# Data Science  - Unit 1 Sprint 3 Module 4

## Module Project: Metrics, Bias and Variance

### Learning Objectives

* Interpret your model results using OLS and Sklearn metrics
* Define and analyze bias in your model
* Define and analyze variance in your model

## Analyzing results from diamonds

Use the seaborn dataset `diamonds` to run a linear regression model and produce the common metrics you would use to evaluate your model's accuracy. 

**Task 1** - Load the data
Load the `diamonds` dataset from the `seaborn` package. 

- Assign the value to an object called `dia`
- Make sure to import the packages you expect to use for an `ols` linear regression model. 

<!-- grade_id: cell-9cb6a33a84a55252, points: 0, grade: false, solution: true -->
```python
#Task 1

#imports



# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-44689b2d030034e2, points: 1, grade: true, solution: false -->
```python
# Task 1 - Tests

assert isinstance(dia, pd.DataFrame)

```

**Task 2** - Conduct EDA on your dataset
- Check for null values. Assign the total number of null values in your dataset to `num_null`

<!-- grade_id: cell-6f304c0a6507dec5, points: 0, grade: false, solution: true -->
```python
#Task 2
# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-fdc8bbdfe5827fa7, points: 1, grade: true, solution: false -->
```python
# Task 2 - Tests

```

**Task 3** - Visualize your feature distributions

- Use seaborn's `pairplot`to visualize the distributions for all your dataset's features.
- You can access the documentation [here](https://seaborn.pydata.org/generated/seaborn.pairplot.html)
- This next task will not be autograded. 

**Task 4**

How would you describe the distribution for the `price` feature?

A: Uniform

B: Right skewed

C: Left skewed

D: Normally distributed

Specify your answer in the next code block using Answer =. For example, if the correct answer is choice B, you'll type Answer = 'B'.

<!-- grade_id: cell-c7981aeaad693e10, points: 0, grade: false, solution: true -->
```python
#Task 4

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-4950fa558630a5ff, points: 1, grade: true, solution: false -->
```python
#Task 4 - Test

```

**Task 5** Check for multicollinearity

- Determine the `pearson` correlations for the `x`, `y`, and `z` columns to `carat`. 
- Assign the value of the correlations to `x_corr`, `y_corr` and `z_corr` respectively. 

<!-- grade_id: cell-3bb28a36542eabdb, points: 0, grade: false, solution: true -->
```python
#Task 5

# YOUR CODE HERE
raise NotImplementedError()


print(x_corr, y_corr, z_corr)
```

<!-- grade_id: cell-a61182842c3afbc9, points: 1, grade: true, solution: false -->
```python
#Task 5 - Test


```

**Task 6** 


Because these three columns share a great deal of correlation with the `carat` feature, it does not make sense to use them as part of our model. Drop the three columns and reassign to the `dia` dataframe. 

<!-- grade_id: cell-9bb880654a9bbb51, points: 0, grade: false, solution: true -->
```python
#Task 6

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b82cf6a65063e5db, points: 1, grade: true, solution: false -->
```python
#Task 6 - Test


```

**Task 7** - OLS Modeling

- Use `carat` as your independent feature. 
- Use the `price` values as your dependent features. 
- Build an OLS model and review the summary report. Make sure to assign a variable called `model`

<!-- grade_id: cell-7cdc6fb8a68ea0a6, points: 0, grade: false, solution: true -->
```python
#Task 7

# YOUR CODE HERE
raise NotImplementedError()

print(model.summary())
```

<!-- grade_id: cell-a286cb026ed74717, points: 1, grade: true, solution: false -->
```python
#Task 7 - Test



```

**Task 8** - Predictions and Residuals 

- Create a new column that includes your model predictions for your features. Name the column `y_pred`
- Calculate the prediction residuals. Assign the values to a column named `residuals`.

<!-- grade_id: cell-194a68304bff654f, points: 0, grade: false, solution: true -->
```python
#Task 8

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-32971dccea2783bd, points: 1, grade: true, solution: false -->
```python
#Task 8 - Test

assert dia.shape == (53940, 9), "Have you created the two columns?"

```

**Task 9** - Metrics

- Determine the values for the **mean absolute error, the mean squared error** and the **root mean squared error** for your previous model. 
- Assign the values as `mae`, `mse`, and `rmse` respectively. 
- *Hint*: We discussed a few methods for this in class. You can refer to this [documentation](https://scikit-learn.org/stable/modules/model_evaluation.html) for other metric values.

<!-- grade_id: cell-8521cfc520c32f52, points: 0, grade: false, solution: true -->
```python
#Task 9

# YOUR CODE HERE
raise NotImplementedError()

print(mae, mse, rmse)
```

<!-- grade_id: cell-320164f97411992e, points: 1, grade: true, solution: false -->
```python
#Task 9 - Test

```

**Task 10** - OLS Modeling, Addtional Features

- Use the `depth`, `table`, and `carat` as your independent features. 
- Use the `price` values as your dependent features. 
- Build an OLS model and review the summary report. Make sure to assign a variable called `model`.  

<!-- grade_id: cell-ae75326bc99156f2, points: 0, grade: false, solution: true -->
```python
#Task 10

# YOUR CODE HERE
raise NotImplementedError()

print(model.summary())
```

<!-- grade_id: cell-30637ae943cc927c, points: 1, grade: true, solution: false -->
```python
#Task 10 - Test


assert len(model.params.index) == 4, "Make sure you've assigned both values."

```

**Task 11** - Predictions and Residuals 

- Create a new column that includes your model predictions for your features. Name the column `y_pred`
- Calculate the residuals. Assign the values to a column named `residuals`.


<!-- grade_id: cell-eda3b235dda0c890, points: 0, grade: false, solution: true -->
```python
#Task 11

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f8dac98f7966de41, points: 1, grade: true, solution: false -->
```python
#Task 11- Test

assert dia.shape == (53940, 9), "Have you created the two columns?"

```

**Task 12** - Metrics

- Determine the values for the **mean absolute error, the mean squared error** and the **root mean squared error** for your previous model. 
- Assign the values as `mae`, `mse`, and `rmse` respectively. 
- *Hint*: We discussed a few methods for this in class. You can refer to this [documentation](https://scikit-learn.org/stable/modules/model_evaluation.html) for other metric values.

<!-- grade_id: cell-29fb7569c52b05de, points: 0, grade: false, solution: true -->
```python
#Task 12

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-9319992c1a2f4f94, points: 1, grade: true, solution: false -->
```python
#Task 12 - Test

```